// variables2.rs
// Make me compile! Execute the command `rustlings hint variables2` if you want a hint :)

// I AM NOT DONE

// use rand::prelude::*;

fn main() {
    let x = 10;
//    let x = rand::thread.rng().gen_range(1,21);
    if x == 10 {
        println!("Ten!");
    } else {
        println!("Not ten!");
    }
}
